# Tutorial 4: Effective Data Visualization 

### Lecture and Tutorial Learning Goals:

Expand your data visualization knowledge and tool set beyond what we have seen and practiced so far. We will move beyond scatter plots and learn other effective ways to visualize data, as well as some general rules of thumb to follow when creating visualizations. All visualization tasks this week will be applied to real world data sets.

After completing this week's lecture and tutorial work, you will be able to:

- Describe when to use the following kinds of visualizations:
    - scatter plots
    - line plots
    - bar plots
- Given a dataset and a question, select from the above plot types to create a visualization that best answers the question
- Given a visualization and a question, evaluate the effectiveness of the visualization and suggest improvements to better answer the question
- Identify rules of thumb for creating effective visualizations
- Define the two key aspects of altair objects:
    - mark objects
    - encodings
- Use the altair library in Python to create and refine the above visualizations using:
    - mark objects: `mark_point`, `mark_line`, `mark_bar`
    - encodings : `x`, `y`, `fill`, `color`, `shape`
    - subplots: `facet`
- Describe the difference in raster and vector output formats
- Use `save` to save visualizations in `png` and `svg` format


Any place you see `___`, you must fill in the function, variable, or data to complete the code. 

```python
### Run this cell before continuing.
import altair as alt
import pandas as pd

alt.data_transformers.disable_max_rows()
```

**Question 0.1** 
<br> {points: 1}

Match the following definitions with the corresponding encoding mapping or function used in Python:

*Definitions*

A. Prevents a chart from being stacked. It preserves the vertical position of a plot while adjusting the horizontal position. 

B. In bar charts, this encoding fills in the bars by a specific colour or separates the counts by a variable different from the x-axis. 

C. In bar charts, it outlines the bars but in scatterplots, it fills in the points (colouring them based on a particular variable aside from the x/y-axis). 

D. This encoding allows further visualization of data by varying data points by shape (modifying their shape based on a particular variable aside from the x/y-axis).

E. Labels the y-axis. 


*Encodings and Functions*

1. `color`
2. `alt.Column`
3. `fill`
4. `alt.Y(title="")`
5. `shape`

_For every description, create an object using the letter associated with the definition and assign it to the corresponding number from the list above. For example: `B = 1`_

<!-- grade_id: cell-b22b3db664eb1abc, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: cell-471b37b1ecc45a3a, points: 1, grade: true, solution: false -->
```python
from hashlib import sha1
assert sha1(str(type(A)).encode("utf-8")+b"aa718d8789417ad8").hexdigest() == "5c9af7469774cb93ea39f0f5603d4e99ec2bfab3", "type of A is not int. Please make sure it is int and not np.int64, etc. You can cast your value into an int using int()"
assert sha1(str(A).encode("utf-8")+b"aa718d8789417ad8").hexdigest() == "7c1b6e40b398b86c2d14cfc8b75794cb17bac77c", "value of A is not correct"

assert sha1(str(type(B)).encode("utf-8")+b"9f8829ed8be1aa44").hexdigest() == "4ac6a39bbb455ea1bc6c04fc0c5ecbeebcc18f49", "type of B is not int. Please make sure it is int and not np.int64, etc. You can cast your value into an int using int()"
assert sha1(str(B).encode("utf-8")+b"9f8829ed8be1aa44").hexdigest() == "6bcdeb4680ca31345f362312b966874bb143fe68", "value of B is not correct"

assert sha1(str(type(C)).encode("utf-8")+b"748eaef662bb9add").hexdigest() == "65bdac93bcdf3c0eb9320b1564adf6dac79f10a7", "type of C is not int. Please make sure it is int and not np.int64, etc. You can cast your value into an int using int()"
assert sha1(str(C).encode("utf-8")+b"748eaef662bb9add").hexdigest() == "1c7f1b599a73e5f358c004f36d84ffea74069c42", "value of C is not correct"

assert sha1(str(type(D)).encode("utf-8")+b"910dff81be557540").hexdigest() == "d5373e7629051e6dd49744bd0ba592816422254f", "type of D is not int. Please make sure it is int and not np.int64, etc. You can cast your value into an int using int()"
assert sha1(str(D).encode("utf-8")+b"910dff81be557540").hexdigest() == "2f9be864e834fc163c20dd0d795a82681f504a4e", "value of D is not correct"

assert sha1(str(type(E)).encode("utf-8")+b"10e4efd31ebe0d79").hexdigest() == "632f17c6eb76432763f87b51da14fa944bc9a577", "type of E is not int. Please make sure it is int and not np.int64, etc. You can cast your value into an int using int()"
assert sha1(str(E).encode("utf-8")+b"10e4efd31ebe0d79").hexdigest() == "c464640eec06368849cf346fa79e18d4d3cf23ad", "value of E is not correct"

print('Success!')
```

**Question 0.2** True or False:
<br> {points: 1}

We should save a plot as an `.svg` file if we want to be able to rescale it without losing quality.

*Assign your answer to an object called `answer0_2`. Make sure your answer is either `True` or `False`.*

<!-- grade_id: cell-8608b643ace9768c, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: cell-dd04803a3b061830, points: 1, grade: true, solution: false -->
```python
from hashlib import sha1
assert sha1(str(type(answer0_2)).encode("utf-8")+b"477fd8254199acd3").hexdigest() == "4e7057f94c39bcfa3009eb5894272208452ba494", "type of answer0_2 is not bool. answer0_2 should be a bool"
assert sha1(str(answer0_2).encode("utf-8")+b"477fd8254199acd3").hexdigest() == "5acf9ae1bffce80d826a997edf025c775b945444", "boolean value of answer0_2 is not correct"

print('Success!')
```

## 1. Data on Personal Medical Costs 

As we saw in the worksheet, data scientists work in all types of organizations and with all kinds of problems. One of these types of organizations are companies in the private sector that work with health data. Today we will be looking at data on personal medical costs. There are varying factors that affect health and consequently medical costs. Our goal for today is to determine how are variables related to the medical costs billed by health insurance companies. 


To analyze this, we will be looking at a dataset that includes the following columns:

* `age`: age of primary beneficiary
* `sex`: insurance contractor gender: female, male
* `bmi`: body mass index, providing an understanding of body, weights that are relatively high or low relative to height, objective index of body weight (kg/$m^{2}$) using the ratio of height to weight, ideally 18.5 to 24.9
* `children`: number of children covered by health insurance / number of dependents
* `smoker`: smoking
* `region`: the beneficiary's residential area in the US: northeast, southeast, southwest, northwest.
* `charges`: individual medical costs billed by health insurance

*This dataset, was taken from the [collection of Data Sets](https://github.com/stedy/Machine-Learning-with-R-datasets) created and curated for the [Machine Learning with R](https://www.packtpub.com/big-data-and-business-intelligence/machine-learning-r) book by Brett Lantz.*

**Question 1.1** Yes or No: 
<br> {points: 1}

Based on the information given in the cell above, do you think the column `charges` includes quantitative/numerical data? 

*Assign your answer to an object called `answer1_1`. Make sure your answer is written in lowercase and is surrounded by quotation marks (e.g. `"yes"` or `"no"`).*

<!-- grade_id: cell-8da0d830fc8f8b0c, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: cell-37a00841b56391ad, points: 1, grade: true, solution: false -->
```python
from hashlib import sha1
assert sha1(str(type(answer1_1)).encode("utf-8")+b"3da4eec3bacb0d18").hexdigest() == "cf6035b282a740825ee3bbeff79ddb86aac9690d", "type of answer1_1 is not str. answer1_1 should be an str"
assert sha1(str(len(answer1_1)).encode("utf-8")+b"3da4eec3bacb0d18").hexdigest() == "1c54a168d74ea897857ceeaf58821a5eb9cf975b", "length of answer1_1 is not correct"
assert sha1(str(answer1_1.lower()).encode("utf-8")+b"3da4eec3bacb0d18").hexdigest() == "4b74cc012a0128ee11773dc50893732f3d671a6f", "value of answer1_1 is not correct"
assert sha1(str(answer1_1).encode("utf-8")+b"3da4eec3bacb0d18").hexdigest() == "4b74cc012a0128ee11773dc50893732f3d671a6f", "correct string value of answer1_1 but incorrect case of letters"

print('Success!')
```

**Question 1.2** Multiple Choice:
<br> {points: 1}

Assuming overplotting is not an issue, which plot would be the most effective to compare the relationship of `age` and `charges`?

A. Scatterplot 

B. Stacked Bar Plot 

C. Bar Plot 

*Assign your answer to an object called `answer1_2`. Make sure your answer is an uppercase letter and is surrounded by quotation marks (e.g. `"F"`).*

<!-- grade_id: cell-776947ebfdd46a40, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: cell-d344b7c03dae8345, points: 1, grade: true, solution: false -->
```python
from hashlib import sha1
assert sha1(str(type(answer1_2)).encode("utf-8")+b"defc916b96141ca1").hexdigest() == "cc4a2d735845aefb77047f72ad27165098da0ad0", "type of answer1_2 is not str. answer1_2 should be an str"
assert sha1(str(len(answer1_2)).encode("utf-8")+b"defc916b96141ca1").hexdigest() == "6d24d2fa20078c66048a7b0cd524c79178f8877e", "length of answer1_2 is not correct"
assert sha1(str(answer1_2.lower()).encode("utf-8")+b"defc916b96141ca1").hexdigest() == "59760d18c5925f89f38df37d6e4c77f12b5969a9", "value of answer1_2 is not correct"
assert sha1(str(answer1_2).encode("utf-8")+b"defc916b96141ca1").hexdigest() == "4cb8c6a790f04c24bfaf5d6df82300344d705c00", "correct string value of answer1_2 but incorrect case of letters"

print('Success!')
```

**Question 1.3**
<br> {points: 1}

Read the `insurance.csv` file in the `data` folder and use `tail` to view the last 6 individuals presented. 

*Assign your answer to an object called `insurance`.*

<!-- grade_id: cell-ace20e7d125da6cc, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
insurance.tail(6)  # preview the last 6 rows of the data set
```

<!-- grade_id: cell-662935f6d4ccb428, points: 1, grade: true, solution: false -->
```python
from hashlib import sha1
assert sha1(str(type(insurance is None)).encode("utf-8")+b"14ac3b7cb76d38cc").hexdigest() == "55b21d79ac4e2871a573f2400cd2c24fead7d8af", "type of insurance is None is not bool. insurance is None should be a bool"
assert sha1(str(insurance is None).encode("utf-8")+b"14ac3b7cb76d38cc").hexdigest() == "89df98ab755ee89a91608581fb11f2e6e0ba8dcf", "boolean value of insurance is None is not correct"

assert sha1(str(type(insurance)).encode("utf-8")+b"1c62da8195639c32").hexdigest() == "bee3e82b2834b4df2c41d50370c8775f76112d63", "type of type(insurance) is not correct"

assert sha1(str(type(insurance.shape)).encode("utf-8")+b"1666389d0f1d7cde").hexdigest() == "e500c7befd143793e01af27099915a36df156778", "type of insurance.shape is not tuple. insurance.shape should be a tuple"
assert sha1(str(len(insurance.shape)).encode("utf-8")+b"1666389d0f1d7cde").hexdigest() == "7ca9cc4b824ef782cccbfb9e70d7a8f3e8a0b1d2", "length of insurance.shape is not correct"
assert sha1(str(sorted(map(str, insurance.shape))).encode("utf-8")+b"1666389d0f1d7cde").hexdigest() == "c6c39cf559902a0a875c6156a30f1c2e03738b2d", "values of insurance.shape are not correct"
assert sha1(str(insurance.shape).encode("utf-8")+b"1666389d0f1d7cde").hexdigest() == "5af964bddb4e25efc5011f97eacecc81494c1212", "order of elements of insurance.shape is not correct"

assert sha1(str(type(round(sum(insurance.age), 2))).encode("utf-8")+b"646afc8f80515627").hexdigest() == "3afd1804726091991e16b3527b2d602fc193039d", "type of round(sum(insurance.age), 2) is not int. Please make sure it is int and not np.int64, etc. You can cast your value into an int using int()"
assert sha1(str(round(sum(insurance.age), 2)).encode("utf-8")+b"646afc8f80515627").hexdigest() == "91f4fb0c5bef4c2ad54eb6257865272b6215e18d", "value of round(sum(insurance.age), 2) is not correct"

assert sha1(str(type(len(insurance.region.unique()))).encode("utf-8")+b"b94c19e885c659c0").hexdigest() == "dfe7b5b2478fd9f0f1b8a260bbb9f41f837ffbbd", "type of len(insurance.region.unique()) is not int. Please make sure it is int and not np.int64, etc. You can cast your value into an int using int()"
assert sha1(str(len(insurance.region.unique())).encode("utf-8")+b"b94c19e885c659c0").hexdigest() == "74bc12878e7593a56ffe00eae43f9e270f4eee0f", "value of len(insurance.region.unique()) is not correct"

print('Success!')
```

**Question 1.4** 
<br> {points: 3}

Looking over the loaded data shown above, what observations can you make about the relationship between medical charges and age? How about medical charges and BMI? Finally, what about medical charges and smoking? 

Also, comment on whether our observations might change if we visualize the data? And/or whether visualizing the data might allow us to more easily make observations about the relationships in the data as opposed to trying to make them directly from the data table?

Answer in the cell below.

<!-- grade_id: cell-62d184813aadd4b3, points: 3, grade: true, solution: true -->
DOUBLE CLICK TO EDIT **THIS CELL** AND REPLACE THIS TEXT WITH YOUR ANSWER.

**Question 1.5**
<br> {points: 1}

According to the National Heart, Lung and Blood Institute of the US: "The higher your BMI, the higher your risk for certain diseases such as heart disease, high blood pressure, type 2 diabetes, gallstones, breathing problems, and certain cancers". 

Based on this information, we can hypothesize that individuals with a higher BMI are likely to have more medical costs. Let's use our data and see if this holds true. Create a scatter plot of `charges` (y-axis) versus `bmi` (x-axis).

In the scaffolding we provide below, we suggest that you set `opacity` to a value between 0.2 and 0.4. `opacity` sets the transparency of points on a scatter plot, and increasing transparencing of points is one tool you can use to deal with over plotting issues.

*Assign your answer to an object called `bmi_plot`. Make sure to label your axes appropriately.*

<!-- grade_id: cell-fa480888e071a7ee, points: 0, grade: false, solution: true -->
```python
# ___ = (
#     alt.Chart(insurance, title=___) # set the title for the entire plot
#     .___(opacity=___)  # Deals with the transparency of the points, set it to an appropiate value
#     .encode(
#         x=alt.X(___),
#         y=alt.Y(___),
#     )
#     .properties(width=380, height=300)  #Remember to set your plot sizes to an appropiate size
# )

# your code here
raise NotImplementedError
bmi_plot
```

<!-- grade_id: cell-b6e79c0645925ade, points: 1, grade: true, solution: false -->
```python
from hashlib import sha1
assert sha1(str(type(bmi_plot is None)).encode("utf-8")+b"c0e2c520fd06e61f").hexdigest() == "f1cc66a88660819986ad2041814d296a13ef2076", "type of bmi_plot is None is not bool. bmi_plot is None should be a bool"
assert sha1(str(bmi_plot is None).encode("utf-8")+b"c0e2c520fd06e61f").hexdigest() == "ad53ded02db096f469089a0ee76972bfeb081094", "boolean value of bmi_plot is None is not correct"

assert sha1(str(type(bmi_plot.encoding.x.field)).encode("utf-8")+b"77e864d2f06bdd83").hexdigest() == "40e505efc5c3e702a52a4c174f583197e56a3fc0", "type of bmi_plot.encoding.x.field is not str. bmi_plot.encoding.x.field should be an str"
assert sha1(str(len(bmi_plot.encoding.x.field)).encode("utf-8")+b"77e864d2f06bdd83").hexdigest() == "2bf9fa732b374235d008a9dfb5595acb98893ed8", "length of bmi_plot.encoding.x.field is not correct"
assert sha1(str(bmi_plot.encoding.x.field.lower()).encode("utf-8")+b"77e864d2f06bdd83").hexdigest() == "c25b366f8fb38a80e5a54c2c33e2f771a9c60788", "value of bmi_plot.encoding.x.field is not correct"
assert sha1(str(bmi_plot.encoding.x.field).encode("utf-8")+b"77e864d2f06bdd83").hexdigest() == "c25b366f8fb38a80e5a54c2c33e2f771a9c60788", "correct string value of bmi_plot.encoding.x.field but incorrect case of letters"

assert sha1(str(type(bmi_plot.encoding.y.field)).encode("utf-8")+b"aa2b9ed9895c6d9d").hexdigest() == "333720a13305058ed18006340734f2cd4d42b540", "type of bmi_plot.encoding.y.field is not str. bmi_plot.encoding.y.field should be an str"
assert sha1(str(len(bmi_plot.encoding.y.field)).encode("utf-8")+b"aa2b9ed9895c6d9d").hexdigest() == "7021b34db5463d4cbbd6b0e47582012338d589fa", "length of bmi_plot.encoding.y.field is not correct"
assert sha1(str(bmi_plot.encoding.y.field.lower()).encode("utf-8")+b"aa2b9ed9895c6d9d").hexdigest() == "c31e5bed2ddd8afa993343881ed9b9d111e5129d", "value of bmi_plot.encoding.y.field is not correct"
assert sha1(str(bmi_plot.encoding.y.field).encode("utf-8")+b"aa2b9ed9895c6d9d").hexdigest() == "c31e5bed2ddd8afa993343881ed9b9d111e5129d", "correct string value of bmi_plot.encoding.y.field but incorrect case of letters"

assert sha1(str(type(bmi_plot.mark)).encode("utf-8")+b"7b8993aade51bd8e").hexdigest() == "fe047e9af34bbbb49a1e110bff3c51a1453f24a7", "type of bmi_plot.mark is not correct"
assert sha1(str(bmi_plot.mark).encode("utf-8")+b"7b8993aade51bd8e").hexdigest() == "f36f0ec7b6b747590dc37c5d66bd7c502703ca6b", "value of bmi_plot.mark is not correct"

assert sha1(str(type(bmi_plot.encoding.x.title != bmi_plot.encoding.x.field)).encode("utf-8")+b"62d65bcf170d1927").hexdigest() == "aff0c14c3115656c0fb4e04fffe7360e6deda46a", "type of bmi_plot.encoding.x.title != bmi_plot.encoding.x.field is not bool. bmi_plot.encoding.x.title != bmi_plot.encoding.x.field should be a bool"
assert sha1(str(bmi_plot.encoding.x.title != bmi_plot.encoding.x.field).encode("utf-8")+b"62d65bcf170d1927").hexdigest() == "5d69c6a6ee77f09a1ae55287812d0ee4a0ce924a", "boolean value of bmi_plot.encoding.x.title != bmi_plot.encoding.x.field is not correct"

assert sha1(str(type(bmi_plot.encoding.y.title != bmi_plot.encoding.y.field)).encode("utf-8")+b"ce7e7c623cfc66a3").hexdigest() == "34e5d3dd2d838f8a7e16f82920a6fc55089c2ed6", "type of bmi_plot.encoding.y.title != bmi_plot.encoding.y.field is not bool. bmi_plot.encoding.y.title != bmi_plot.encoding.y.field should be a bool"
assert sha1(str(bmi_plot.encoding.y.title != bmi_plot.encoding.y.field).encode("utf-8")+b"ce7e7c623cfc66a3").hexdigest() == "f8635aea2e7c5b7085b3f691d9e1076f6435ad84", "boolean value of bmi_plot.encoding.y.title != bmi_plot.encoding.y.field is not correct"

print('Success!')
```

**Question 1.6**
<br> {points: 3}

Analysis: Comment on the effectiveness of the plot. Take into consideration the rules of thumb discussed in lecture. Also comment on what could be improved for this plot and also what is done correctly. 

Answer in the cell below.

<!-- grade_id: cell-a3673af9bcf8288b, points: 3, grade: true, solution: true -->
DOUBLE CLICK TO EDIT **THIS CELL** AND REPLACE THIS TEXT WITH YOUR ANSWER.

**Question 1.7**
<br> {points: 3}

Analysis: What do you observe from the scatter plot? Do the data suggest that there might be evidence of a relationship between BMI and medical costs of individuals? 
From this plot alone, can we say higher BMI causes higher medical charges? Why or why not? 

Answer in the cell below. 

<!-- grade_id: cell-93c1d5cb7d1bcbea, points: 3, grade: true, solution: true -->
DOUBLE CLICK TO EDIT **THIS CELL** AND REPLACE THIS TEXT WITH YOUR ANSWER.

**Question 1.8**
<br> {points: 3}

Again, based on information from the National Heart, Lung and Blood Institute of the US, smoking cigarettes is said to be a risk factor for obesity. Create the same plot as you did in **Question 1.5** but this time add the `color` encoding to observe if smoking might affect the body mass of individuals. You may want to pass `opacity = 0.4` to the scatter geometric object to make the scatter points translucent (just for your own ease of visualization; you don't have to and we won't check that when grading).

*Assign your answer to an object called `smoke_plot`. Make sure to label your axes appropriately.*

<!-- grade_id: cell-2128ed1aa06e5391, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
smoke_plot
```

<!-- grade_id: cell-a5dc720f9e5813bd, points: 3, grade: true, solution: false -->
```python
from hashlib import sha1
assert sha1(str(type(smoke_plot is None)).encode("utf-8")+b"4e8a2347a99fd111").hexdigest() == "a16f8c658c383d3b549ec525674d7fdceedf37ca", "type of smoke_plot is None is not bool. smoke_plot is None should be a bool"
assert sha1(str(smoke_plot is None).encode("utf-8")+b"4e8a2347a99fd111").hexdigest() == "573509636511666f838f27abb6a04b4c99a05844", "boolean value of smoke_plot is None is not correct"

assert sha1(str(type(smoke_plot.encoding.x.field)).encode("utf-8")+b"04b91e5f7779734f").hexdigest() == "69c1c3b7c705f0417d366af5041c13aa0e235aa2", "type of smoke_plot.encoding.x.field is not str. smoke_plot.encoding.x.field should be an str"
assert sha1(str(len(smoke_plot.encoding.x.field)).encode("utf-8")+b"04b91e5f7779734f").hexdigest() == "6ed1f05f340461bb479ec3cb0cf49f1e7babef01", "length of smoke_plot.encoding.x.field is not correct"
assert sha1(str(smoke_plot.encoding.x.field.lower()).encode("utf-8")+b"04b91e5f7779734f").hexdigest() == "963e3762701fed4b46499d564ae7d37c60fa7a10", "value of smoke_plot.encoding.x.field is not correct"
assert sha1(str(smoke_plot.encoding.x.field).encode("utf-8")+b"04b91e5f7779734f").hexdigest() == "963e3762701fed4b46499d564ae7d37c60fa7a10", "correct string value of smoke_plot.encoding.x.field but incorrect case of letters"

assert sha1(str(type(smoke_plot.encoding.y.field)).encode("utf-8")+b"f832156d0a7f9193").hexdigest() == "b9323da3037c42b4e01ae4694a3666b9f531df5a", "type of smoke_plot.encoding.y.field is not str. smoke_plot.encoding.y.field should be an str"
assert sha1(str(len(smoke_plot.encoding.y.field)).encode("utf-8")+b"f832156d0a7f9193").hexdigest() == "7bfee464f922b24dd8d691301e7274d70ca91660", "length of smoke_plot.encoding.y.field is not correct"
assert sha1(str(smoke_plot.encoding.y.field.lower()).encode("utf-8")+b"f832156d0a7f9193").hexdigest() == "1d6c69ccbb6f4127152bc19bdbcda087ed8de4ed", "value of smoke_plot.encoding.y.field is not correct"
assert sha1(str(smoke_plot.encoding.y.field).encode("utf-8")+b"f832156d0a7f9193").hexdigest() == "1d6c69ccbb6f4127152bc19bdbcda087ed8de4ed", "correct string value of smoke_plot.encoding.y.field but incorrect case of letters"

assert sha1(str(type(smoke_plot.encoding.color.field)).encode("utf-8")+b"f3a9e309a9dfbcda").hexdigest() == "5b3eae788daa5c9c47a3dedbdc23214a97ace4e9", "type of smoke_plot.encoding.color.field is not str. smoke_plot.encoding.color.field should be an str"
assert sha1(str(len(smoke_plot.encoding.color.field)).encode("utf-8")+b"f3a9e309a9dfbcda").hexdigest() == "6cc589c2b27e3b91d46ad3f17cea6ffc295c375f", "length of smoke_plot.encoding.color.field is not correct"
assert sha1(str(smoke_plot.encoding.color.field.lower()).encode("utf-8")+b"f3a9e309a9dfbcda").hexdigest() == "4e12db7fcdeed48a2923c8995a990c0651dd1584", "value of smoke_plot.encoding.color.field is not correct"
assert sha1(str(smoke_plot.encoding.color.field).encode("utf-8")+b"f3a9e309a9dfbcda").hexdigest() == "4e12db7fcdeed48a2923c8995a990c0651dd1584", "correct string value of smoke_plot.encoding.color.field but incorrect case of letters"

assert sha1(str(type(smoke_plot.mark.type)).encode("utf-8")+b"709639564908751b").hexdigest() == "636031857a000295d50adb756ef97d0b494e1ebf", "type of smoke_plot.mark.type is not str. smoke_plot.mark.type should be an str"
assert sha1(str(len(smoke_plot.mark.type)).encode("utf-8")+b"709639564908751b").hexdigest() == "c2aa8af981a8ed55c1c8590ae9fa81cc0e6eac81", "length of smoke_plot.mark.type is not correct"
assert sha1(str(smoke_plot.mark.type.lower()).encode("utf-8")+b"709639564908751b").hexdigest() == "1969d4d8107fa2133c6792020ae1bb363e11720b", "value of smoke_plot.mark.type is not correct"
assert sha1(str(smoke_plot.mark.type).encode("utf-8")+b"709639564908751b").hexdigest() == "1969d4d8107fa2133c6792020ae1bb363e11720b", "correct string value of smoke_plot.mark.type but incorrect case of letters"

assert sha1(str(type(smoke_plot.encoding.x.title != smoke_plot.encoding.x.field)).encode("utf-8")+b"5d229565efe90204").hexdigest() == "5af0035b794ff88cfef9fe98f08456ee4a3214c6", "type of smoke_plot.encoding.x.title != smoke_plot.encoding.x.field is not bool. smoke_plot.encoding.x.title != smoke_plot.encoding.x.field should be a bool"
assert sha1(str(smoke_plot.encoding.x.title != smoke_plot.encoding.x.field).encode("utf-8")+b"5d229565efe90204").hexdigest() == "499ef404c523bd0355f22b39f141f48ade8609a3", "boolean value of smoke_plot.encoding.x.title != smoke_plot.encoding.x.field is not correct"

assert sha1(str(type(smoke_plot.encoding.y.title != smoke_plot.encoding.y.field)).encode("utf-8")+b"d2cb5f531cdb12ed").hexdigest() == "712dc0ac48c2ef62a307273ec27d9421fe88c542", "type of smoke_plot.encoding.y.title != smoke_plot.encoding.y.field is not bool. smoke_plot.encoding.y.title != smoke_plot.encoding.y.field should be a bool"
assert sha1(str(smoke_plot.encoding.y.title != smoke_plot.encoding.y.field).encode("utf-8")+b"d2cb5f531cdb12ed").hexdigest() == "b95ef974144586344b25e7de44cea0e9426cddfa", "boolean value of smoke_plot.encoding.y.title != smoke_plot.encoding.y.field is not correct"

assert sha1(str(type(smoke_plot.encoding.color.title != smoke_plot.encoding.color.field)).encode("utf-8")+b"cc5c11fc990a82e5").hexdigest() == "ee99d51b01f0f09f6021693b29aa7581384a2731", "type of smoke_plot.encoding.color.title != smoke_plot.encoding.color.field is not bool. smoke_plot.encoding.color.title != smoke_plot.encoding.color.field should be a bool"
assert sha1(str(smoke_plot.encoding.color.title != smoke_plot.encoding.color.field).encode("utf-8")+b"cc5c11fc990a82e5").hexdigest() == "75f7109322d31c532dbb15ce79cfd8643f1668ab", "boolean value of smoke_plot.encoding.color.title != smoke_plot.encoding.color.field is not correct"

print('Success!')
```

**Question 1.9.0** (Analyzing the Graph) True or False: 
<br> {points: 1}

Smokers generally have a lower BMI than non-smokers. 

*Assign your answer to an object called `answer1_9_0`. Make sure your answer is either `True` or `False`.*

<!-- grade_id: cell-4302d377b0fc06e3, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: cell-9d85f368775a121b, points: 1, grade: true, solution: false -->
```python
from hashlib import sha1
assert sha1(str(type(answer1_9_0)).encode("utf-8")+b"8b921544d3eb966d").hexdigest() == "644ba59a1eb140b304e587ec34a65c6875ae67c2", "type of answer1_9_0 is not bool. answer1_9_0 should be a bool"
assert sha1(str(answer1_9_0).encode("utf-8")+b"8b921544d3eb966d").hexdigest() == "f226fcf948959b1c77dac7daaa0eaa1af401998f", "boolean value of answer1_9_0 is not correct"

print('Success!')
```

**Question 1.9.1** (Analyzing the Graph) True or False: 
<br> {points: 1}

Smokers generally have higher medical charges than non-smokers.

*Assign your answer to an object called `answer1_9_1`.Make sure your answer is either `True` or `False`.*

<!-- grade_id: cell-b9f9c54000889b5c, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: cell-b82577eedde24cb1, points: 1, grade: true, solution: false -->
```python
from hashlib import sha1
assert sha1(str(type(answer1_9_1)).encode("utf-8")+b"782c933c8e3389bd").hexdigest() == "fa7f83b998462201c8fa705de24d8792ce031365", "type of answer1_9_1 is not bool. answer1_9_1 should be a bool"
assert sha1(str(answer1_9_1).encode("utf-8")+b"782c933c8e3389bd").hexdigest() == "0c13bb78f8858fdfb6183c66d2a5f850e146d5c1", "boolean value of answer1_9_1 is not correct"

print('Success!')
```

**Question 1.10**
<br> {points: 1}

Finally, create a bar graph that displays the proportion of smokers for both females and males in the data set. Use sex as the horizontal axis, and colour the bars to differentiate between smokers / nonsmokers. This could, for example, be used help us determine whether we should consider smoking behaviour when exploring whether there is a relationship between sex and medical costs.

*Assign your answer to an object called `bar_plot`. Make sure to label your axes appropriately.*

>*Note - many historical datasets treated sex as a variable where the possible values are only binary: male or female. This representation in this question reflects how the data were historically collected and is not meant to imply that we believe that sex is binary.*

<!-- grade_id: cell-19a1b424460c544e, points: 0, grade: false, solution: true -->
```python
# ___ = (
#     alt.Chart(insurance, title=___)
#     .___()
#     .encode(
#         y=alt.Y(
#             "count(___)",
#             stack="normalize",
#             axis=alt.Axis(format="%"),
#             title=___,
#         ),
#         x=alt.X(___, title=___),
#         fill=alt.Fill(___, title="Does the person smoke?"),
#     )
# )

# your code here
raise NotImplementedError
bar_plot
```

<!-- grade_id: cell-37fe64d61b7366b4, points: 1, grade: true, solution: false -->
```python
from hashlib import sha1
assert sha1(str(type(bar_plot is None)).encode("utf-8")+b"dd23011d43702646").hexdigest() == "3566edb71a2b42c0d410f19b20d6ebf4c64bd07e", "type of bar_plot is None is not bool. bar_plot is None should be a bool"
assert sha1(str(bar_plot is None).encode("utf-8")+b"dd23011d43702646").hexdigest() == "b8027de5ac713d5c164b681848fa17438736672d", "boolean value of bar_plot is None is not correct"

assert sha1(str(type(bar_plot.encoding.x.field)).encode("utf-8")+b"c4b38bf36d7c3588").hexdigest() == "11ed03928b672e19fccc1e66d9f7f0f0a9d7e0da", "type of bar_plot.encoding.x.field is not str. bar_plot.encoding.x.field should be an str"
assert sha1(str(len(bar_plot.encoding.x.field)).encode("utf-8")+b"c4b38bf36d7c3588").hexdigest() == "fda9c86fc90ee45f0c64228035061728e3e081a1", "length of bar_plot.encoding.x.field is not correct"
assert sha1(str(bar_plot.encoding.x.field.lower()).encode("utf-8")+b"c4b38bf36d7c3588").hexdigest() == "b3bc2f024aedacf088ce055d4e767a297f124de7", "value of bar_plot.encoding.x.field is not correct"
assert sha1(str(bar_plot.encoding.x.field).encode("utf-8")+b"c4b38bf36d7c3588").hexdigest() == "b3bc2f024aedacf088ce055d4e767a297f124de7", "correct string value of bar_plot.encoding.x.field but incorrect case of letters"

assert sha1(str(type(bar_plot.encoding.fill.field)).encode("utf-8")+b"51311cc1e374810b").hexdigest() == "116e55c35dce9ea2c858608207d6b31263b9afe7", "type of bar_plot.encoding.fill.field is not str. bar_plot.encoding.fill.field should be an str"
assert sha1(str(len(bar_plot.encoding.fill.field)).encode("utf-8")+b"51311cc1e374810b").hexdigest() == "ff78da8a4b453bdfc50521c9c926dcc4ce010905", "length of bar_plot.encoding.fill.field is not correct"
assert sha1(str(bar_plot.encoding.fill.field.lower()).encode("utf-8")+b"51311cc1e374810b").hexdigest() == "bcea5228630a05983659f95eedb638b82a65a0c4", "value of bar_plot.encoding.fill.field is not correct"
assert sha1(str(bar_plot.encoding.fill.field).encode("utf-8")+b"51311cc1e374810b").hexdigest() == "bcea5228630a05983659f95eedb638b82a65a0c4", "correct string value of bar_plot.encoding.fill.field but incorrect case of letters"

assert sha1(str(type(bar_plot.encoding.y.stack)).encode("utf-8")+b"d4a37cc535bcd3b9").hexdigest() == "6c159ce803e1de13f516b545cb99bb162185d306", "type of bar_plot.encoding.y.stack is not str. bar_plot.encoding.y.stack should be an str"
assert sha1(str(len(bar_plot.encoding.y.stack)).encode("utf-8")+b"d4a37cc535bcd3b9").hexdigest() == "b633ac42bd05cf6b935b3690b69543e9d15f14fc", "length of bar_plot.encoding.y.stack is not correct"
assert sha1(str(bar_plot.encoding.y.stack.lower()).encode("utf-8")+b"d4a37cc535bcd3b9").hexdigest() == "138ab561d4932a8864b0e3e43d009a6f8098529f", "value of bar_plot.encoding.y.stack is not correct"
assert sha1(str(bar_plot.encoding.y.stack).encode("utf-8")+b"d4a37cc535bcd3b9").hexdigest() == "138ab561d4932a8864b0e3e43d009a6f8098529f", "correct string value of bar_plot.encoding.y.stack but incorrect case of letters"

assert sha1(str(type(bar_plot.mark)).encode("utf-8")+b"9c999925b51d5e9c").hexdigest() == "bae0e6ff0839884659b4a05bd34874639011f667", "type of bar_plot.mark is not str. bar_plot.mark should be an str"
assert sha1(str(len(bar_plot.mark)).encode("utf-8")+b"9c999925b51d5e9c").hexdigest() == "0d61333dcf9d911e570ec73e511970cd6436b315", "length of bar_plot.mark is not correct"
assert sha1(str(bar_plot.mark.lower()).encode("utf-8")+b"9c999925b51d5e9c").hexdigest() == "52698064b9151fb8fd8e677e91c448c004d0a375", "value of bar_plot.mark is not correct"
assert sha1(str(bar_plot.mark).encode("utf-8")+b"9c999925b51d5e9c").hexdigest() == "52698064b9151fb8fd8e677e91c448c004d0a375", "correct string value of bar_plot.mark but incorrect case of letters"

assert sha1(str(type(bar_plot.encoding.fill.title != bar_plot.encoding.fill.field)).encode("utf-8")+b"66e783d0458e2410").hexdigest() == "cf4cc85613ad70a77ec422b872311dce446eb3d4", "type of bar_plot.encoding.fill.title != bar_plot.encoding.fill.field is not bool. bar_plot.encoding.fill.title != bar_plot.encoding.fill.field should be a bool"
assert sha1(str(bar_plot.encoding.fill.title != bar_plot.encoding.fill.field).encode("utf-8")+b"66e783d0458e2410").hexdigest() == "20f8c11523d49c8734e5ca513785b4700fa98eae", "boolean value of bar_plot.encoding.fill.title != bar_plot.encoding.fill.field is not correct"

assert sha1(str(type(bar_plot.encoding.x.title != bar_plot.encoding.x.field)).encode("utf-8")+b"a68a482fc719ea84").hexdigest() == "72c383dc3f6f08a8b45ad4ea9aa1b764be087065", "type of bar_plot.encoding.x.title != bar_plot.encoding.x.field is not bool. bar_plot.encoding.x.title != bar_plot.encoding.x.field should be a bool"
assert sha1(str(bar_plot.encoding.x.title != bar_plot.encoding.x.field).encode("utf-8")+b"a68a482fc719ea84").hexdigest() == "5330cc223277781175c96c1f7c96ec78b82b0adf", "boolean value of bar_plot.encoding.x.title != bar_plot.encoding.x.field is not correct"

assert sha1(str(type(bar_plot.encoding.y.title != bar_plot.encoding.y.field)).encode("utf-8")+b"018f2e87c94066fe").hexdigest() == "417e84c886033c33d076f2242820e6840beb2aad", "type of bar_plot.encoding.y.title != bar_plot.encoding.y.field is not bool. bar_plot.encoding.y.title != bar_plot.encoding.y.field should be a bool"
assert sha1(str(bar_plot.encoding.y.title != bar_plot.encoding.y.field).encode("utf-8")+b"018f2e87c94066fe").hexdigest() == "4e68d570dc511ff0494ae0b2b566b7758be6a34c", "boolean value of bar_plot.encoding.y.title != bar_plot.encoding.y.field is not correct"

print('Success!')
```

**Question 1.11**
<br> {points: 1}

Based on the graph, is the proportion of smokers higher amongst men or women?

*Assign your answer to an object called `answer1_11`. Make sure your answer is in lowercase and is surrounded by quotation marks (e.g. `"male"` or `"female"`).*

<!-- grade_id: cell-123d8e2e2b0e9555, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: cell-204b484044f109f1, points: 1, grade: true, solution: false -->
```python
from hashlib import sha1
assert sha1(str(type(answer1_11)).encode("utf-8")+b"c07c5e8158fffd67").hexdigest() == "f1b18de1915cc5284720e38a77a88184ac84cd20", "type of answer1_11 is not str. answer1_11 should be an str"
assert sha1(str(len(answer1_11)).encode("utf-8")+b"c07c5e8158fffd67").hexdigest() == "77de673d485fc0b5e1610bf495a651e787fd0c5e", "length of answer1_11 is not correct"
assert sha1(str(answer1_11.lower()).encode("utf-8")+b"c07c5e8158fffd67").hexdigest() == "7d95ea1d3369f025450be6186fd9b17612c63538", "value of answer1_11 is not correct"
assert sha1(str(answer1_11).encode("utf-8")+b"c07c5e8158fffd67").hexdigest() == "7d95ea1d3369f025450be6186fd9b17612c63538", "correct string value of answer1_11 but incorrect case of letters"

print('Success!')
```

## 2. Color Palettes (beyond the defaults)
{points: 3}

In the worksheet and this tutorial, you have seen the same colours again and again. These are from the default `altair` color palette. What if you want different colors? There are other color palettes available in Altair and there are two ways to change the color palette (examples are provided below as well):

- You can change the colormap (or colorscheme) by specifying its name as a string to scheme inside `alt.Scale`. 

All the available colormaps can be viewed on [this page](https://vega.github.io/vega/docs/schemes/), which also lists what type of data the colormap is useful for (categorical, sequential, diverging, cyclic). The default categorical colormap used in Altair is "Tableau10", which consists of 10 colors and starts with a blue, orange, and red color. You can also specify the color by its [names](https://stackoverflow.com/questions/22408237/named-colors-in-matplotlib/37232760#37232760).


```python
cars = pd.read_csv("data/cars.csv")
cars.head()
```

```python
cars_plot = (
    alt.Chart(cars)
    .mark_bar()
    .encode(
        y=alt.Y(
            "count(Cylinders)",
            stack="normalize",
            axis=alt.Axis(format="%"),
            title="Proportion",
        ),
        x=alt.X("Origin"),
        fill=alt.Fill(
            "Cylinders:N",
        ),  # N specifies that this is a "nominal" variable; e.g. a categorical variable
    )
    .properties(width=380, height=300)
    .configure_axis(labelFontSize=18, titleFontSize=18)
    .configure_title(fontSize=20)
)

cars_plot
```

```python
# change the color palette in scale
cars_plot = (
    alt.Chart(cars)
    .mark_bar()
    .encode(
        y=alt.Y(
            "count(Cylinders)",
            stack="normalize",
            axis=alt.Axis(format="%"),
            title="Proportion",
        ),
        x=alt.X("Origin"),
        fill=alt.Fill(
            "Cylinders:N", scale=alt.Scale(scheme="dark2")
        ),  # N is to specify this is a categorical variable
    )
    .properties(width=380, height=300)
    .configure_axis(labelFontSize=18, titleFontSize=18)
    .configure_title(fontSize=20)
)

cars_plot
```

Please change the color palette to `accent`. *Assign your answer to an object called `cars_plot`.*

<!-- grade_id: cell-54bd5139136b22a6, points: 3, grade: true, solution: true -->
```python
# your code here
raise NotImplementedError
cars_plot
```

## 3. Fast-Food Chains in the United States (Continued)
<br> {points: 6}

In Worksheet 04, we explored this data set through some visualizations. Now, it is is all up to you. The goal of this assignment is to create **one** plot that can help you figure out which restaurant to open and where! After that you need to write a paragraph explaining your visualization and why you chose it. Also, explain your conclusion from the visualization and reasoning as to how you came to that conclusion. If you need to bring in outside information to help you answer your question, please feel free to do so.  Finally, if there is some way that you could improve your visualization, but don't yet know how to do it, please explain what you would do if you knew how.

In answering this question, there is no need to restrict yourself to the west coast of the USA. Consider all states that you have data for. You have a variety of graphs to choose from, but before starting the assignment, discuss with a partner which plot would be the most optimal to answer this question.

<img src="mcdonalds.jpg" width = "300"/>




<!-- grade_id: cell-86b3015872de2602, points: 3, grade: true, solution: true -->
```python
# write the code for your plot here
# your code here
raise NotImplementedError
```

<!-- grade_id: cell-6ac8b7c5b2701b78, points: 3, grade: true, solution: true -->
*Write a paragraph explaining your visualization and why you chose it. Also explain your conclusion from the visualization and reasoning as to how you came to that conclusion. If you need to bring in outside information to help you answer your question, please feel free to do so. Finally, if there is some way that you could improve your visualization, but don't yet know how to do it, please explain what you would do if you knew how.*

DOUBLE CLICK TO EDIT **THIS CELL** AND REPLACE THIS TEXT WITH YOUR ANSWER.