# Worksheet 12 - The Bootstrap

We will cover 2 things today:

1. [Estimation when you only have 1 sample (like in real life!)](#Estimation-with-only-one-sample)
2. [Getting Python + Jupyter working for you outside of our course](#Getting-Python-+-Jupyter-working-for-you-outside-of-our-course)

### Lecture and Tutorial Learning Goals:

After completing this week's lecture and tutorial work, you will be able to:
- Explain why we don't have a sampling distribution in practice/real life.
- Define bootstrapping.
- Use Python to create a bootstrap distribution to approximate a sampling distribution.
- Contrast bootstrap and sampling distributions.

```python
### Run this cell before continuing.
import altair as alt
import numpy as np
import pandas as pd
from sklearn.utils import resample

alt.data_transformers.disable_max_rows()
```

**Question 1.1** True/False:
<br> {points: 1}

In real life, we typically take many samples from the population and create a sampling distribution when we perform estimation. True or false?

*Assign your answer to an object called `answer1_1`. Your answer should be a boolean. i.e. `True` or `False`.*

<!-- grade_id: cell-03cc711477e4baf7, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: cell-e8aa24da1ef5f535, points: 1, grade: true, solution: false -->
```python
from hashlib import sha1
assert sha1(str(type(answer1_1)).encode("utf-8")+b"35e13fc434714847").hexdigest() == "d7b702950a8f5c3b75658bb704424882db1b4481", "type of answer1_1 is not bool. answer1_1 should be a bool"
assert sha1(str(answer1_1).encode("utf-8")+b"35e13fc434714847").hexdigest() == "d88a0f03113d63bfd811a037dd0e910c1d39addd", "boolean value of answer1_1 is not correct"

print('Success!')
```

**Question 1.2** Ordering
<br> {points: 1}

Correctly re-order the steps for creating a bootstrap sample from those listed below. 

1. record the observation's value
2. repeat the above the same number of times as there are observations in the original sample 
3. return the observation to the original sample
4. randomly draw an observation from the original sample (which was drawn from the population)

Create your answer by reordering values below in the `answer1_2` vector with the correct order for the steps above for creating a bootstrap sample.

<!-- grade_id: cell-dbdb7facfc6bda43, points: 0, grade: false, solution: true -->
```python
answer1_2 = [1, 2, 3, 4]  # reorder the values!
# your code here
raise NotImplementedError
```

<!-- grade_id: cell-8a257d7535ce46fe, points: 1, grade: true, solution: false -->
```python
from hashlib import sha1
assert sha1(str(type("".join(map(str, answer1_2)))).encode("utf-8")+b"eabf544d426997ee").hexdigest() == "9d81ee77c601ca2af164965195c78e599ac8ee98", "type of \"\".join(map(str, answer1_2)) is not str. \"\".join(map(str, answer1_2)) should be an str"
assert sha1(str(len("".join(map(str, answer1_2)))).encode("utf-8")+b"eabf544d426997ee").hexdigest() == "8402f2e7f79595acd8fcb248a648e4aade21eb5d", "length of \"\".join(map(str, answer1_2)) is not correct"
assert sha1(str("".join(map(str, answer1_2)).lower()).encode("utf-8")+b"eabf544d426997ee").hexdigest() == "df0831d389c9f7abbd52f87ed4e09b86ea128efd", "value of \"\".join(map(str, answer1_2)) is not correct"
assert sha1(str("".join(map(str, answer1_2))).encode("utf-8")+b"eabf544d426997ee").hexdigest() == "df0831d389c9f7abbd52f87ed4e09b86ea128efd", "correct string value of \"\".join(map(str, answer1_2)) but incorrect case of letters"

print('Success!')
```

**Question 1.3** Multiple choice
<br> {points: 1}

From the list below, choose the correct description of a bootstrap distribution for a point estimate:

A. a list of point estimates calculated from many samples drawn with replacement from the population

B. a list of point estimates calculated from many samples drawn without replacement from the population

C. a list of point estimates calculated from bootstrap samples drawn with replacement from a single sample (that was drawn from the population)

D. a list of point estimates calculated from bootstrap samples drawn without replacement from a single sample (that was drawn from the population)

*Assign your answer to an object called `answer1_3`. Your answer should be an uppercase letter and is surrounded by quotes. (e.g. `"F"`)*

<!-- grade_id: cell-fe8959defed4a439, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: cell-59dbb6a427a133c5, points: 1, grade: true, solution: false -->
```python
from hashlib import sha1
assert sha1(str(type(answer1_3)).encode("utf-8")+b"09d0d7f98d3f9991").hexdigest() == "4316a3d60889a4b4232d7127220c2ec5b80aca6e", "type of answer1_3 is not str. answer1_3 should be an str"
assert sha1(str(len(answer1_3)).encode("utf-8")+b"09d0d7f98d3f9991").hexdigest() == "ad6a5d9ebd81ceed099d08ab06e69f0b569bfd8c", "length of answer1_3 is not correct"
assert sha1(str(answer1_3.lower()).encode("utf-8")+b"09d0d7f98d3f9991").hexdigest() == "4273550d2cf18e4a585b1dc51631352895d4f7da", "value of answer1_3 is not correct"
assert sha1(str(answer1_3).encode("utf-8")+b"09d0d7f98d3f9991").hexdigest() == "5be543b327cce9edecce35f2bc4d6cf0e73755ad", "correct string value of answer1_3 but incorrect case of letters"

print('Success!')
```

**Question 1.4** Multiple choice
<br> {points: 1}

From the list below, choose the correct explanation of why, when performing estimation, we want to report a plausible **range** for the true population quantity we are trying to estimate along with the point estimate:

A. The point estimate is our best guess at the true population quantity we are trying to estimate

B. The point estimate will often not be the exact value of the true population quantity we are trying to estimate

C. The value of a point estimate from one sample might very well be different than the value of a point estimate from another sample.

D. B & C

F. A & C

E. None of the above

*Assign your answer to an object called `answer1_4`. Your answer should be an uppercase letter and is surrounded by quotes (e.g. `"F"`).*

<!-- grade_id: cell-fd7e89338bbb71ca, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: cell-9448b3731eccbe8c, points: 1, grade: true, solution: false -->
```python
from hashlib import sha1
assert sha1(str(type(answer1_4)).encode("utf-8")+b"4335fa2975b463d2").hexdigest() == "abc6301e3c02eb9500c98a68c67ed110793905e2", "type of answer1_4 is not str. answer1_4 should be an str"
assert sha1(str(len(answer1_4)).encode("utf-8")+b"4335fa2975b463d2").hexdigest() == "79759c872df3e17f0eb58817321bb7141f428b7f", "length of answer1_4 is not correct"
assert sha1(str(answer1_4.lower()).encode("utf-8")+b"4335fa2975b463d2").hexdigest() == "79cbb64faa0cf14124a1d75c66c8e9f82587ff9e", "value of answer1_4 is not correct"
assert sha1(str(answer1_4).encode("utf-8")+b"4335fa2975b463d2").hexdigest() == "23dbf5624f5c93e847d7f5842c773dc761055c59", "correct string value of answer1_4 but incorrect case of letters"

print('Success!')
```

###  Continuing with our virtual population of Canadian seniors from last worksheet

Here we re-create the virtual population (ages of all Canadian seniors) we used in the last worksheet. It was bounded by realistic values ($\geq$ 65 and $\leq$ 118):

```python
# run this cell to simulate a finite population
np.random.seed(4321)
can_seniors = pd.DataFrame({"age": (np.random.exponential(1 / 0.1, 2000000)) ** 2 + 65})
can_seniors = can_seniors.query("65<=age<=118")
can_seniors.head()
```

Let's remind ourselves of what this population looks like:

```python
# run this cell
pop_dist = (
    alt.Chart(can_seniors, title="Population distribution")
    .mark_bar()
    .encode(x=alt.X("age:Q", title="Age (years)", bin=alt.Bin(maxbins=30)), y="count()")
    .properties(width=400, height=400)
    .configure_title(fontSize=20)
    .configure_axis(labelFontSize=18, titleFontSize=15)
)
pop_dist
```

### Estimate the mean age of Canadian Seniors

Let's say we are interested in estimating the mean age of Canadian Seniors. Given that we have the population (we created it) we could just calculate the mean age from this population data. However in real life, we usually only have one small-ish sample from the population. Also, from our experimentation with sampling distributions, we know that different random samples will give us different point estimates. We also know from these experiments that the point estimates from different random samples will mostly be close to the true population quanitity we are trying to estimate, and how close depends on the sample size.

What about in real life though, when we only have one sample? Can we say how close? Or at least give some plausible range of where we would expect the population quanitity we are trying to estimate to fall? Yes! We can do this using a method called bootstrapping! Let's explore how to create a bootstrap distribution from a single sample using Python and then we will discuss how the bootstrap distribution relates to the sampling distribution, and what it can tell us about the true population quantity we are trying to estimate.

Let's draw a single sample of size 40 from the population and visualize it:

```python
# run this cell
one_sample = can_seniors.sample(40, random_state=12345)
one_sample.head()
```

```python
# run this cell
one_sample_dist = (
    alt.Chart(one_sample, title="Distribution of one sample")
    .mark_bar()
    .encode(x=alt.X("age:Q", title="Age (years)", bin=alt.Bin(step=1)), y="count()")
)
one_sample_dist
```

**Question 1.5** 
<br> {points: 1}

Use `mean` to calculate the mean age (our point estimate of interest) from the random sample you just took (`one_sample`):

Name this data frame `one_sample_estimates` which has a single column named `mean`.

<!-- grade_id: cell-443a7d1dcaf35e9a, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
one_sample_estimates
```

<!-- grade_id: cell-8e9e83895e2967dd, points: 1, grade: true, solution: false -->
```python
from hashlib import sha1
assert sha1(str(type(one_sample_estimates.shape[0])).encode("utf-8")+b"bbd50c8a593854fb").hexdigest() == "5b037a949a473c23e7efaa3d9a36d4289056c79a", "type of one_sample_estimates.shape[0] is not int. Please make sure it is int and not np.int64, etc. You can cast your value into an int using int()"
assert sha1(str(one_sample_estimates.shape[0]).encode("utf-8")+b"bbd50c8a593854fb").hexdigest() == "f0e625df37e6e861db511487658529dfcd336e6d", "value of one_sample_estimates.shape[0] is not correct"

assert sha1(str(type(one_sample_estimates.shape[1])).encode("utf-8")+b"42d5f20be4f7c2d4").hexdigest() == "7756d52c684df416028f746ef9300f98f558e715", "type of one_sample_estimates.shape[1] is not int. Please make sure it is int and not np.int64, etc. You can cast your value into an int using int()"
assert sha1(str(one_sample_estimates.shape[1]).encode("utf-8")+b"42d5f20be4f7c2d4").hexdigest() == "509f6fc0715e3fd2e9f19c5c2595931161c3d3a9", "value of one_sample_estimates.shape[1] is not correct"

assert sha1(str(type("".join(one_sample_estimates.columns))).encode("utf-8")+b"47de7c8d2ba370c2").hexdigest() == "067e37400691717bc1e92ec0633eacd7cd93a30e", "type of \"\".join(one_sample_estimates.columns) is not str. \"\".join(one_sample_estimates.columns) should be an str"
assert sha1(str(len("".join(one_sample_estimates.columns))).encode("utf-8")+b"47de7c8d2ba370c2").hexdigest() == "45689d220cb675e808177375c1662f340276f992", "length of \"\".join(one_sample_estimates.columns) is not correct"
assert sha1(str("".join(one_sample_estimates.columns).lower()).encode("utf-8")+b"47de7c8d2ba370c2").hexdigest() == "a9f37442d5289550b66e52d6db6fa304808a9220", "value of \"\".join(one_sample_estimates.columns) is not correct"
assert sha1(str("".join(one_sample_estimates.columns)).encode("utf-8")+b"47de7c8d2ba370c2").hexdigest() == "a9f37442d5289550b66e52d6db6fa304808a9220", "correct string value of \"\".join(one_sample_estimates.columns) but incorrect case of letters"

assert sha1(str(type(round(one_sample_estimates['mean'], 2))).encode("utf-8")+b"a878b9569b13aac7").hexdigest() == "b7be10c6e927624f3f6195d0a88a1ba0ff26df5d", "type of round(one_sample_estimates['mean'], 2) is not correct"
assert sha1(str(round(one_sample_estimates['mean'], 2)).encode("utf-8")+b"a878b9569b13aac7").hexdigest() == "1513d4820217d60c14b3ff162c781f954cccff54", "value of round(one_sample_estimates['mean'], 2) is not correct"

print('Success!')
```

**Question 1.6** 
<br> {points: 1}

To generate a single bootstrap sample in Python, we can use the `resample` function from `sklearn` package, and set the `replace` argument to its default value of `True`.

Use `resample` to take a single bootstrap sample from the sample you drew from the population. Name this bootstrap sample `boot1`.

Use 4321 as the `random_state`.

<!-- grade_id: cell-239c20c4566f1f3c, points: 0, grade: false, solution: true -->
```python
# ___ = resample(___, n_samples=___, replace=___, random_state=___)

# your code here
raise NotImplementedError
boot1.head()
```

<!-- grade_id: cell-af716b8959aee1ac, points: 1, grade: true, solution: false -->
```python
from hashlib import sha1
assert sha1(str(type("".join(boot1.columns))).encode("utf-8")+b"60c73aa993afc7a0").hexdigest() == "ab7bf96cff33f13e77fad1f3ddee6c0f2738e25f", "type of \"\".join(boot1.columns) is not str. \"\".join(boot1.columns) should be an str"
assert sha1(str(len("".join(boot1.columns))).encode("utf-8")+b"60c73aa993afc7a0").hexdigest() == "077df13b345ff022a2ccf31256d595c51c0fae6e", "length of \"\".join(boot1.columns) is not correct"
assert sha1(str("".join(boot1.columns).lower()).encode("utf-8")+b"60c73aa993afc7a0").hexdigest() == "c4200ee02d9720d3da452557cc8433f48d4c99cb", "value of \"\".join(boot1.columns) is not correct"
assert sha1(str("".join(boot1.columns)).encode("utf-8")+b"60c73aa993afc7a0").hexdigest() == "c4200ee02d9720d3da452557cc8433f48d4c99cb", "correct string value of \"\".join(boot1.columns) but incorrect case of letters"

assert sha1(str(type(boot1.shape[0])).encode("utf-8")+b"4f61e902b27deccb").hexdigest() == "45d7ae1f5cc652d1cacdd904fde8dd1e3028326d", "type of boot1.shape[0] is not int. Please make sure it is int and not np.int64, etc. You can cast your value into an int using int()"
assert sha1(str(boot1.shape[0]).encode("utf-8")+b"4f61e902b27deccb").hexdigest() == "503d90073d91e331bc9124f22614147667cc2c77", "value of boot1.shape[0] is not correct"

assert sha1(str(type(round(sum(boot1.age), 2))).encode("utf-8")+b"0dbcba323aa51bd4").hexdigest() == "7d92da55fe0ea335fbac40607073345dfc20001b", "type of round(sum(boot1.age), 2) is not float. Please make sure it is float and not np.float64, etc. You can cast your value into a float using float()"
assert sha1(str(round(round(sum(boot1.age), 2), 2)).encode("utf-8")+b"0dbcba323aa51bd4").hexdigest() == "cdc515b02f048b6c2e60e8e690551a843f27e792", "value of round(sum(boot1.age), 2) is not correct (rounded to 2 decimal places)"

assert sha1(str(type(sum(range(boot1.shape[0])))).encode("utf-8")+b"eecb144117c8dc8b").hexdigest() == "b3ca78bc58cd1450ad2369ed53154cd81059da6d", "type of sum(range(boot1.shape[0])) is not int. Please make sure it is int and not np.int64, etc. You can cast your value into an int using int()"
assert sha1(str(sum(range(boot1.shape[0]))).encode("utf-8")+b"eecb144117c8dc8b").hexdigest() == "ca91bc2366ef7d72c9bd58048e32ee30354a8274", "value of sum(range(boot1.shape[0])) is not correct"

print('Success!')
```

**Question 1.7** Multiple choice
<br> {points: 1}

Why do we change `replace` to `TRUE`?

A. Taking a bootstrap sample involves drawing observations from the original population without replacement

B. Taking a bootstrap sample involves drawing observations from the original population with replacement

C. Taking a bootstrap sample involves drawing observations from the original sample without replacement

D. Taking a bootstrap sample involves drawing observations from the original sample with replacement

*Assign your answer to an object called `answer1_7`. Your answer should be an uppercase letter and is surrounded by quotes (e.g. `"F"`).*

<!-- grade_id: cell-ba91c2de68609fe7, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: cell-1b79e9b905b70215, points: 1, grade: true, solution: false -->
```python
from hashlib import sha1
assert sha1(str(type(answer1_7)).encode("utf-8")+b"8ef6cc3157515455").hexdigest() == "cc3e388146a61cb58067543c53f4030ca32a5da7", "type of answer1_7 is not str. answer1_7 should be an str"
assert sha1(str(len(answer1_7)).encode("utf-8")+b"8ef6cc3157515455").hexdigest() == "0f8f46eb96dae87ae51f98245e34fe6879aa5559", "length of answer1_7 is not correct"
assert sha1(str(answer1_7.lower()).encode("utf-8")+b"8ef6cc3157515455").hexdigest() == "dffd8a5bb3c4194ef0c3d9fe03c0be3eb1023f3a", "value of answer1_7 is not correct"
assert sha1(str(answer1_7).encode("utf-8")+b"8ef6cc3157515455").hexdigest() == "fc8b885d020bfc3f248d92ad8b17babd665eaa49", "correct string value of answer1_7 but incorrect case of letters"

print('Success!')
```

**Question 1.8** 
<br> {points: 1}

Visualize the distribution of the bootstrap sample you just took (`boot1`) that was just created by plotting a histogram using `mark_bar` and `step = 1` in the `alt.Bin` argument in the `y` encoding. Name the plot `boot1_dist` and give the plot and the x-axis a descriptive label.

<!-- grade_id: cell-389271affd1fc5af, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
boot1_dist
```

<!-- grade_id: cell-3b7b4b4c4cc4971e, points: 1, grade: true, solution: false -->
```python
from hashlib import sha1
assert sha1(str(type(boot1_dist.encoding.x.field)).encode("utf-8")+b"e4b2f0938ac1a41b").hexdigest() == "12c057a11788f004d163c1a59cdbe1a57a314228", "type of boot1_dist.encoding.x.field is not str. boot1_dist.encoding.x.field should be an str"
assert sha1(str(len(boot1_dist.encoding.x.field)).encode("utf-8")+b"e4b2f0938ac1a41b").hexdigest() == "3ff05a256e7feb05e90a0b6d9fca17ad1aece11a", "length of boot1_dist.encoding.x.field is not correct"
assert sha1(str(boot1_dist.encoding.x.field.lower()).encode("utf-8")+b"e4b2f0938ac1a41b").hexdigest() == "4e00ed626734bc58958553d4aee96bde9fafa8d3", "value of boot1_dist.encoding.x.field is not correct"
assert sha1(str(boot1_dist.encoding.x.field).encode("utf-8")+b"e4b2f0938ac1a41b").hexdigest() == "4e00ed626734bc58958553d4aee96bde9fafa8d3", "correct string value of boot1_dist.encoding.x.field but incorrect case of letters"

assert sha1(str(type(boot1_dist.mark)).encode("utf-8")+b"033820b56dbf9057").hexdigest() == "626727b85625a1cf8ccf291769647d326699371e", "type of boot1_dist.mark is not str. boot1_dist.mark should be an str"
assert sha1(str(len(boot1_dist.mark)).encode("utf-8")+b"033820b56dbf9057").hexdigest() == "9582dbbdbb96cba2c800b374eacf2861974d3533", "length of boot1_dist.mark is not correct"
assert sha1(str(boot1_dist.mark.lower()).encode("utf-8")+b"033820b56dbf9057").hexdigest() == "e250db812a0935692141eb203ead6c8e01f6e16a", "value of boot1_dist.mark is not correct"
assert sha1(str(boot1_dist.mark).encode("utf-8")+b"033820b56dbf9057").hexdigest() == "e250db812a0935692141eb203ead6c8e01f6e16a", "correct string value of boot1_dist.mark but incorrect case of letters"

assert sha1(str(type(boot1_dist.data.shape[0])).encode("utf-8")+b"02e5484d778decad").hexdigest() == "e3f7908b594749a2085abf8bfa2937ef060042b2", "type of boot1_dist.data.shape[0] is not int. Please make sure it is int and not np.int64, etc. You can cast your value into an int using int()"
assert sha1(str(boot1_dist.data.shape[0]).encode("utf-8")+b"02e5484d778decad").hexdigest() == "e1e575734167507acc518d2199c7fde0a4d8596e", "value of boot1_dist.data.shape[0] is not correct"

assert sha1(str(type(round(boot1_dist.data.sum(), 2))).encode("utf-8")+b"503204094a2a28ca").hexdigest() == "406384ba9d2094fa8b9735e0fc17aa6ac22839b8", "type of round(boot1_dist.data.sum(), 2) is not correct"
assert sha1(str(round(boot1_dist.data.sum(), 2)).encode("utf-8")+b"503204094a2a28ca").hexdigest() == "398c8ffeda3d18ba7fb04391c483e3904ac5ce1c", "value of round(boot1_dist.data.sum(), 2) is not correct"

assert sha1(str(type(boot1_dist.encoding.x.field != boot1_dist.encoding.x.title)).encode("utf-8")+b"921d19f0d591a526").hexdigest() == "79901a8990749fd128b07bce145bddab2164dc3b", "type of boot1_dist.encoding.x.field != boot1_dist.encoding.x.title is not bool. boot1_dist.encoding.x.field != boot1_dist.encoding.x.title should be a bool"
assert sha1(str(boot1_dist.encoding.x.field != boot1_dist.encoding.x.title).encode("utf-8")+b"921d19f0d591a526").hexdigest() == "8c07edc1a849b695641a257835117981bba38def", "boolean value of boot1_dist.encoding.x.field != boot1_dist.encoding.x.title is not correct"

assert sha1(str(type(boot1_dist.title is not None)).encode("utf-8")+b"411298342186b8a8").hexdigest() == "97dfc51a161d920fc71346527ac829155137331b", "type of boot1_dist.title is not None is not bool. boot1_dist.title is not None should be a bool"
assert sha1(str(boot1_dist.title is not None).encode("utf-8")+b"411298342186b8a8").hexdigest() == "9dd3c62b0b651e6ac7920e51bcd00850d01d70d3", "boolean value of boot1_dist.title is not None is not correct"

print('Success!')
```

Let's now compare our bootstrap sample to the original random sample that we drew from the population:

```python
# run this code cell
one_sample_dist
```

```python
# run this code cell
boot1_dist
```

Earlier we calculate the mean of our original sample to be about 79.6 years. What is the mean of our bootstrap sample?

```python
# run this cell
pd.DataFrame({"mean": boot1.mean()})
```

We see that original sample distrbution and the bootstrap sample distribution are of similar shape, but not identical. They also have different means. The difference of the frequency of the values in the bootstrap sample (and the difference of the value of the mean) comes from sampling from the original sample with replacement. Why sample with replacement? If we didn't we would end up with the original sample again. What we are trying to do with bootstrapping is to mimic drawing another sample from the population, without actually doing that. 

Why are we doing this? As mentioned earlier, in real life we typically only have one sample and thus we cannot create a sampling distribution that we can use to tell us about how we might expected our point estimate to behave if we took another sample. So what we can do, is use our sample as an estimate of our population, and sample from that with replacement (i.e., bootstrapping) many times to create many bootstrap samples. We can then calculate point estimates for each bootstrap sample and create a bootstrap distribution of our point estimates and use this as a proxy for a sampling distribution. We can finally use this bootstrap distribution of our point estimates to suggest how we might expected our point estimate to behave if we took another sample.

**Question 1.9** 
<br> {points: 1}

What do 6 different bootstrap samples look like? Use `resample` to create a single data frame with 6 bootstrap samples drawn from the original sample we drew from the population, `one_sample`. Assign a new column called `replicate` to mark the sample number `(from 0 to 5)`. Name the data frame `boot6`.

Set the seed as 1234.

<!-- grade_id: cell-6607b2dec92a2913, points: 0, grade: false, solution: true -->
```python
np.random.seed(1234)  # DO NOT CHANGE!
# your code here
raise NotImplementedError
print(boot6.head())
print(boot6.tail())
```

<!-- grade_id: cell-e57717e45df34f9a, points: 1, grade: true, solution: false -->
```python
from hashlib import sha1
assert sha1(str(type("".join(boot6.columns))).encode("utf-8")+b"19e6a6bc38ab444c").hexdigest() == "aaffa57cdcbe650b50a29483a2f878978bf4947f", "type of \"\".join(boot6.columns) is not str. \"\".join(boot6.columns) should be an str"
assert sha1(str(len("".join(boot6.columns))).encode("utf-8")+b"19e6a6bc38ab444c").hexdigest() == "a161ed67597b2c3083a5cebd859e82d8d312f7bb", "length of \"\".join(boot6.columns) is not correct"
assert sha1(str("".join(boot6.columns).lower()).encode("utf-8")+b"19e6a6bc38ab444c").hexdigest() == "355afa5ece676047b07dacfbc154251f2e9ade24", "value of \"\".join(boot6.columns) is not correct"
assert sha1(str("".join(boot6.columns)).encode("utf-8")+b"19e6a6bc38ab444c").hexdigest() == "355afa5ece676047b07dacfbc154251f2e9ade24", "correct string value of \"\".join(boot6.columns) but incorrect case of letters"

assert sha1(str(type(boot6.shape[0])).encode("utf-8")+b"44f750aa85a46c70").hexdigest() == "9625e1212fb8877f68e6e9eda567f9a022d7f64c", "type of boot6.shape[0] is not int. Please make sure it is int and not np.int64, etc. You can cast your value into an int using int()"
assert sha1(str(boot6.shape[0]).encode("utf-8")+b"44f750aa85a46c70").hexdigest() == "6ba62d9a6d140f98f6811f3cc5d9d2c878d02b41", "value of boot6.shape[0] is not correct"

assert sha1(str(type(round(sum(boot6.age), 2))).encode("utf-8")+b"ed99bf28a454df28").hexdigest() == "444086b2510fe35e496fe7af1887e132f0145b44", "type of round(sum(boot6.age), 2) is not float. Please make sure it is float and not np.float64, etc. You can cast your value into a float using float()"
assert sha1(str(round(round(sum(boot6.age), 2), 2)).encode("utf-8")+b"ed99bf28a454df28").hexdigest() == "a5e4f22b5a7827195fe4d457840f2084797afee7", "value of round(sum(boot6.age), 2) is not correct (rounded to 2 decimal places)"

assert sha1(str(type(sum(boot6.replicate.unique()))).encode("utf-8")+b"157d522ca75ca9d0").hexdigest() == "01bf3de05fd66909af8c4c8115c6b52d4384b629", "type of sum(boot6.replicate.unique()) is not correct"
assert sha1(str(sum(boot6.replicate.unique())).encode("utf-8")+b"157d522ca75ca9d0").hexdigest() == "9fac426953420a3551e9d1490e1e526b5b0b975d", "value of sum(boot6.replicate.unique()) is not correct"

print('Success!')
```

**Question 2.0** 
<br> {points: 1}

Now visualize the six bootstrap sample distributions from `boot6` by faceting the `replicate` column. Name the plot object `boot6_dist` and give the plot and the x-axis a descriptive label.

<!-- grade_id: cell-d1d36a4806776c6c, points: 0, grade: false, solution: true -->
```python
# ___ = (
#     alt.Chart(___, title=___)
#     .___()
#     .encode(x=alt.X(___, title=___, bin=alt.Bin(step=1)), y=___)
#     .facet(___, columns=3)
# )
# your code here
raise NotImplementedError
boot6_dist
```

<!-- grade_id: cell-38e6f9e544c5ce17, points: 1, grade: true, solution: false -->
```python
from hashlib import sha1
assert sha1(str(type(boot6_dist.facet.field)).encode("utf-8")+b"30c2b479f4ee4167").hexdigest() == "51e4e0a15322b5e2628c75c4c6ba05230c20a3d2", "type of boot6_dist.facet.field is not str. boot6_dist.facet.field should be an str"
assert sha1(str(len(boot6_dist.facet.field)).encode("utf-8")+b"30c2b479f4ee4167").hexdigest() == "a7c078ab875b671c996e394f6d85eb0170bcefc7", "length of boot6_dist.facet.field is not correct"
assert sha1(str(boot6_dist.facet.field.lower()).encode("utf-8")+b"30c2b479f4ee4167").hexdigest() == "5a611c71c41399b4f0b6f106eaa00f5e590cc40d", "value of boot6_dist.facet.field is not correct"
assert sha1(str(boot6_dist.facet.field).encode("utf-8")+b"30c2b479f4ee4167").hexdigest() == "5a611c71c41399b4f0b6f106eaa00f5e590cc40d", "correct string value of boot6_dist.facet.field but incorrect case of letters"

assert sha1(str(type(boot6_dist.data.shape[0])).encode("utf-8")+b"e8aa71452aabedfc").hexdigest() == "06dfe4ed2174f7a6fe5d152f7fb8fbdf3dbc438b", "type of boot6_dist.data.shape[0] is not int. Please make sure it is int and not np.int64, etc. You can cast your value into an int using int()"
assert sha1(str(boot6_dist.data.shape[0]).encode("utf-8")+b"e8aa71452aabedfc").hexdigest() == "eb7ae4029f109fe2d29b3e9f4ab7323aec47e3e9", "value of boot6_dist.data.shape[0] is not correct"

assert sha1(str(type(boot6_dist.title is not None)).encode("utf-8")+b"ca8e589deba087c5").hexdigest() == "a98b7390e4ebecfd878e0084d40b14893ccb9a4b", "type of boot6_dist.title is not None is not bool. boot6_dist.title is not None should be a bool"
assert sha1(str(boot6_dist.title is not None).encode("utf-8")+b"ca8e589deba087c5").hexdigest() == "57a805e966c60b55f72da1edcdf9cfcaeecfe616", "boolean value of boot6_dist.title is not None is not correct"

print('Success!')
```

**Question 2.1** 
<br> {points: 1}

Calculate the mean of these 6 bootstrap samples using `groupby` and `mean` and save result into a column called `mean`. The data frame should have two columns `replicate` and `mean`. Name the data frame `boot6_means`.

<!-- grade_id: cell-b9692d80805c724a, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
boot6_means
```

<!-- grade_id: cell-f2cfe994fece602d, points: 1, grade: true, solution: false -->
```python
from hashlib import sha1
assert sha1(str(type(boot6_means.shape[0])).encode("utf-8")+b"0470f483acc4b724").hexdigest() == "144ea8c669643c4ea3d7a11490e2c1dee3bb1d37", "type of boot6_means.shape[0] is not int. Please make sure it is int and not np.int64, etc. You can cast your value into an int using int()"
assert sha1(str(boot6_means.shape[0]).encode("utf-8")+b"0470f483acc4b724").hexdigest() == "b188be414f996a9cdc8f234bb09304dd1df70618", "value of boot6_means.shape[0] is not correct"

assert sha1(str(type(boot6_means.shape[1])).encode("utf-8")+b"1d6007273063fcb9").hexdigest() == "929a3546d83b12f767d9dae8af0c23379726ef20", "type of boot6_means.shape[1] is not int. Please make sure it is int and not np.int64, etc. You can cast your value into an int using int()"
assert sha1(str(boot6_means.shape[1]).encode("utf-8")+b"1d6007273063fcb9").hexdigest() == "66298568376ce71a8336f044b1a6f0f442c095cb", "value of boot6_means.shape[1] is not correct"

assert sha1(str(type("".join(boot6_means.columns))).encode("utf-8")+b"0870ed23b657fca6").hexdigest() == "5b65c1df1371fc4b85e8638d620147de77c2dd52", "type of \"\".join(boot6_means.columns) is not str. \"\".join(boot6_means.columns) should be an str"
assert sha1(str(len("".join(boot6_means.columns))).encode("utf-8")+b"0870ed23b657fca6").hexdigest() == "ee69c9946e1a14be843c5e89727cdacdd5770b6d", "length of \"\".join(boot6_means.columns) is not correct"
assert sha1(str("".join(boot6_means.columns).lower()).encode("utf-8")+b"0870ed23b657fca6").hexdigest() == "654a544b54ca8d8c5656d77ceba2cefcd2b0bda3", "value of \"\".join(boot6_means.columns) is not correct"
assert sha1(str("".join(boot6_means.columns)).encode("utf-8")+b"0870ed23b657fca6").hexdigest() == "654a544b54ca8d8c5656d77ceba2cefcd2b0bda3", "correct string value of \"\".join(boot6_means.columns) but incorrect case of letters"

assert sha1(str(type(round(boot6_means["mean"][0], 2))).encode("utf-8")+b"56959642ed881b08").hexdigest() == "e4a997b950c476c7c2b2cc4b107d2818e51f359f", "type of round(boot6_means[\"mean\"][0], 2) is not correct"
assert sha1(str(round(boot6_means["mean"][0], 2)).encode("utf-8")+b"56959642ed881b08").hexdigest() == "189ecae86e863186c8bee6f439d4b01250d97437", "value of round(boot6_means[\"mean\"][0], 2) is not correct"

print('Success!')
```

**Question 2.2** 
<br> {points: 1}

Let's now take 1000 bootstrap samples from the original sample we drew from the population (`one_sample`) using `resample`. Assign a new column called `replicate` to mark the sample number `(from 0 to 99)`. Name the data frame `boot1000`.

Set the seed as 1234.

<!-- grade_id: cell-cd89ca0c9e456e1b, points: 0, grade: false, solution: true -->
```python
np.random.seed(1234)  # DO NOT CHANGE!
# your code here
raise NotImplementedError
print(boot1000.head())
print(boot1000.tail())
```

<!-- grade_id: cell-343d880b6f801822, points: 1, grade: true, solution: false -->
```python
from hashlib import sha1
assert sha1(str(type("".join(boot1000.columns))).encode("utf-8")+b"0cdd386aea9e0528").hexdigest() == "cd72b903a2703f0b556acfe51c2c81e0d10753ef", "type of \"\".join(boot1000.columns) is not str. \"\".join(boot1000.columns) should be an str"
assert sha1(str(len("".join(boot1000.columns))).encode("utf-8")+b"0cdd386aea9e0528").hexdigest() == "b7d93a818aba98e8dc98fe8540fe93428467882a", "length of \"\".join(boot1000.columns) is not correct"
assert sha1(str("".join(boot1000.columns).lower()).encode("utf-8")+b"0cdd386aea9e0528").hexdigest() == "35b52af9b1da4d2728c2dd52997553d713e47cd5", "value of \"\".join(boot1000.columns) is not correct"
assert sha1(str("".join(boot1000.columns)).encode("utf-8")+b"0cdd386aea9e0528").hexdigest() == "35b52af9b1da4d2728c2dd52997553d713e47cd5", "correct string value of \"\".join(boot1000.columns) but incorrect case of letters"

assert sha1(str(type(boot1000.shape[0])).encode("utf-8")+b"f4fb29092f656dea").hexdigest() == "0591a308cc217d7d135ab7a97428c560a31e169b", "type of boot1000.shape[0] is not int. Please make sure it is int and not np.int64, etc. You can cast your value into an int using int()"
assert sha1(str(boot1000.shape[0]).encode("utf-8")+b"f4fb29092f656dea").hexdigest() == "a6fc6d674f3031f058171aff87fcf7431125df7f", "value of boot1000.shape[0] is not correct"

assert sha1(str(type(round(sum(boot1000.age), 2))).encode("utf-8")+b"554940391e9be41f").hexdigest() == "b3d495d08a1abc00fcffc291dbadf105765c3513", "type of round(sum(boot1000.age), 2) is not float. Please make sure it is float and not np.float64, etc. You can cast your value into a float using float()"
assert sha1(str(round(round(sum(boot1000.age), 2), 2)).encode("utf-8")+b"554940391e9be41f").hexdigest() == "086ccc642d03121a465444c0673a23f606581d92", "value of round(sum(boot1000.age), 2) is not correct (rounded to 2 decimal places)"

assert sha1(str(type(sum([x for x in range(boot1000.shape[0])]))).encode("utf-8")+b"c34723b524c47d80").hexdigest() == "e2405a7a5deb254984655bf43df92b89eb2104af", "type of sum([x for x in range(boot1000.shape[0])]) is not int. Please make sure it is int and not np.int64, etc. You can cast your value into an int using int()"
assert sha1(str(sum([x for x in range(boot1000.shape[0])])).encode("utf-8")+b"c34723b524c47d80").hexdigest() == "6e2db136c17fd7a79b959bb795a96169b3e2400c", "value of sum([x for x in range(boot1000.shape[0])]) is not correct"

print('Success!')
```

**Question 2.3** 
<br> {points: 1}

Calculate the mean of these 1000 bootstrap samples using `groupby` and `mean` and save the result into a column called `mean`. Name the data frame `boot1000_means`.

<!-- grade_id: cell-51d5f2afbb677369, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
print(boot1000_means.head())
print(boot1000_means.tail())
```

<!-- grade_id: cell-bcf16788d57cef52, points: 1, grade: true, solution: false -->
```python
from hashlib import sha1
assert sha1(str(type(boot1000_means.shape[0])).encode("utf-8")+b"84878aa99093ddd8").hexdigest() == "f7f33c3798c287b40fed05cb4c18a8b6a5c4bb52", "type of boot1000_means.shape[0] is not int. Please make sure it is int and not np.int64, etc. You can cast your value into an int using int()"
assert sha1(str(boot1000_means.shape[0]).encode("utf-8")+b"84878aa99093ddd8").hexdigest() == "6e83151b1b7970c9d5c5e91c9b5e137d99d6b079", "value of boot1000_means.shape[0] is not correct"

assert sha1(str(type(boot1000_means.shape[1])).encode("utf-8")+b"201833a7208a1517").hexdigest() == "b4662c5daa142aefd358fa94469894081e423960", "type of boot1000_means.shape[1] is not int. Please make sure it is int and not np.int64, etc. You can cast your value into an int using int()"
assert sha1(str(boot1000_means.shape[1]).encode("utf-8")+b"201833a7208a1517").hexdigest() == "72eaa0662b3ec6e7fea890db14c14f38db41c21a", "value of boot1000_means.shape[1] is not correct"

assert sha1(str(type("".join(boot1000_means.columns))).encode("utf-8")+b"04999d75e0a3f8c5").hexdigest() == "4f175e2ed72fb805bfb620dd2319efe747797a44", "type of \"\".join(boot1000_means.columns) is not str. \"\".join(boot1000_means.columns) should be an str"
assert sha1(str(len("".join(boot1000_means.columns))).encode("utf-8")+b"04999d75e0a3f8c5").hexdigest() == "8de2e6f6e951a31bed1e0596bdc58ba83924919b", "length of \"\".join(boot1000_means.columns) is not correct"
assert sha1(str("".join(boot1000_means.columns).lower()).encode("utf-8")+b"04999d75e0a3f8c5").hexdigest() == "2dface7cc837581ea88390c9cc9c347cd694464d", "value of \"\".join(boot1000_means.columns) is not correct"
assert sha1(str("".join(boot1000_means.columns)).encode("utf-8")+b"04999d75e0a3f8c5").hexdigest() == "2dface7cc837581ea88390c9cc9c347cd694464d", "correct string value of \"\".join(boot1000_means.columns) but incorrect case of letters"

assert sha1(str(type(round(boot1000_means["mean"][0], 2))).encode("utf-8")+b"521cb29d1f9655ef").hexdigest() == "f3bf0d014a0d807f5590f3db5c3cae9a9d1fc942", "type of round(boot1000_means[\"mean\"][0], 2) is not correct"
assert sha1(str(round(boot1000_means["mean"][0], 2)).encode("utf-8")+b"521cb29d1f9655ef").hexdigest() == "c44bb12728b93aa77b3b29407435e2b48efa5032", "value of round(boot1000_means[\"mean\"][0], 2) is not correct"

print('Success!')
```

**Question 2.4** 
<br> {points: 1}

Visualize the distribution of the bootstrap sample point estimates (`boot1000_means`) you just calculated by plotting a histogram using `mark_bar` and `step = 1` in the `alt.Bin` argument in the `y` encoding. Name the plot `boot_est_dist` and give the plot and the x-axis a descriptive label.

<!-- grade_id: cell-21f093a752cc5d8c, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
boot_est_dist
```

<!-- grade_id: cell-13a4f6d6fbcd8cbb, points: 1, grade: true, solution: false -->
```python
from hashlib import sha1
assert sha1(str(type(boot_est_dist.encoding.x.field)).encode("utf-8")+b"25585e7cb6b02d14").hexdigest() == "8e7a28348e2c081708ad9cf6acc5e23280efdcbd", "type of boot_est_dist.encoding.x.field is not str. boot_est_dist.encoding.x.field should be an str"
assert sha1(str(len(boot_est_dist.encoding.x.field)).encode("utf-8")+b"25585e7cb6b02d14").hexdigest() == "a6702497c029992ba13daf642d8e8bf1af8ce452", "length of boot_est_dist.encoding.x.field is not correct"
assert sha1(str(boot_est_dist.encoding.x.field.lower()).encode("utf-8")+b"25585e7cb6b02d14").hexdigest() == "13cbf55ba275cf8d9c90c57ae8aade77121951fd", "value of boot_est_dist.encoding.x.field is not correct"
assert sha1(str(boot_est_dist.encoding.x.field).encode("utf-8")+b"25585e7cb6b02d14").hexdigest() == "13cbf55ba275cf8d9c90c57ae8aade77121951fd", "correct string value of boot_est_dist.encoding.x.field but incorrect case of letters"

assert sha1(str(type(boot_est_dist.mark)).encode("utf-8")+b"fb2aa06a2d708472").hexdigest() == "9a6750a81d5553403b1b23e3cb37e544534e6550", "type of boot_est_dist.mark is not str. boot_est_dist.mark should be an str"
assert sha1(str(len(boot_est_dist.mark)).encode("utf-8")+b"fb2aa06a2d708472").hexdigest() == "7b9961552d4b82a0ee71d193d31b62ed16593844", "length of boot_est_dist.mark is not correct"
assert sha1(str(boot_est_dist.mark.lower()).encode("utf-8")+b"fb2aa06a2d708472").hexdigest() == "a68718e864bde5d787740d5bad561d88ce82976d", "value of boot_est_dist.mark is not correct"
assert sha1(str(boot_est_dist.mark).encode("utf-8")+b"fb2aa06a2d708472").hexdigest() == "a68718e864bde5d787740d5bad561d88ce82976d", "correct string value of boot_est_dist.mark but incorrect case of letters"

assert sha1(str(type(boot_est_dist.data.shape[0])).encode("utf-8")+b"3e7f705dbb9b7591").hexdigest() == "c20fb3a5cbf4e9e9dd6110f7fb6829f7ad363461", "type of boot_est_dist.data.shape[0] is not int. Please make sure it is int and not np.int64, etc. You can cast your value into an int using int()"
assert sha1(str(boot_est_dist.data.shape[0]).encode("utf-8")+b"3e7f705dbb9b7591").hexdigest() == "f2e1aabe67975e763dafe978e9d5b1c726053f01", "value of boot_est_dist.data.shape[0] is not correct"

assert sha1(str(type(round(sum(boot_est_dist.data.sum()), 2))).encode("utf-8")+b"5a7112f068f344e3").hexdigest() == "93c848922c01d59945e523168d56901b383b294a", "type of round(sum(boot_est_dist.data.sum()), 2) is not float. Please make sure it is float and not np.float64, etc. You can cast your value into a float using float()"
assert sha1(str(round(round(sum(boot_est_dist.data.sum()), 2), 2)).encode("utf-8")+b"5a7112f068f344e3").hexdigest() == "b8df91c23f5329d3e5ed921d96fd074edd5c7371", "value of round(sum(boot_est_dist.data.sum()), 2) is not correct (rounded to 2 decimal places)"

assert sha1(str(type(boot_est_dist.encoding.x.field != boot_est_dist.encoding.x.title)).encode("utf-8")+b"e9cde656476d4a4c").hexdigest() == "5d8a3c964b808f7bdb3a68ea98ab6339ed638ab7", "type of boot_est_dist.encoding.x.field != boot_est_dist.encoding.x.title is not bool. boot_est_dist.encoding.x.field != boot_est_dist.encoding.x.title should be a bool"
assert sha1(str(boot_est_dist.encoding.x.field != boot_est_dist.encoding.x.title).encode("utf-8")+b"e9cde656476d4a4c").hexdigest() == "08c11995780756a54b28fab055c4590f6d0dba2f", "boolean value of boot_est_dist.encoding.x.field != boot_est_dist.encoding.x.title is not correct"

assert sha1(str(type(boot_est_dist.title is not None)).encode("utf-8")+b"4d66ea8037ca87f2").hexdigest() == "42d1ba7e301ab5a929764243a90c7e8b534348a7", "type of boot_est_dist.title is not None is not bool. boot_est_dist.title is not None should be a bool"
assert sha1(str(boot_est_dist.title is not None).encode("utf-8")+b"4d66ea8037ca87f2").hexdigest() == "c654d5a7f2973fbe0bce4653924614525cc15036", "boolean value of boot_est_dist.title is not None is not correct"

print('Success!')
```

How does the bootstrap distribution above compare to the sampling distribution? Let's visualize them side by side:

```python
# run this cell

# create sampling distribution histogram
np.random.seed(4321)
samples = []
for rep in range(1500):
    sample = resample(can_seniors, n_samples=40, replace=True)
    sample = sample.assign(replicate=rep)
    samples.append(sample)
samples = pd.concat([samples[i] for i in range(len(samples))])

sample_estimates = (
    samples.groupby("replicate").mean().reset_index().rename(columns={"age": "mean"})
)
sampling_dist = (
    alt.Chart(sample_estimates, title="Sampling distribution of the sample means")
    .mark_bar()
    .encode(
        x=alt.X("mean:Q", title="Mean Age (years)", bin=alt.Bin(step=1)), y="count()"
    )
) + alt.Chart({"values": [{"x": 85, "y": 200}]}).mark_text(
    text="mean = " + str(round(np.mean(sample_estimates["mean"]), 2))
).encode(
    x="x:Q", y="y:Q"
)

# annotate distribution mean to the bootstrap dist
boot_est_dist = boot_est_dist + alt.Chart({"values": [{"x": 90, "y": 140}]}).mark_text(
    text="mean = " + str(round(np.mean(boot1000_means["mean"]), 2))
).encode(x="x:Q", y="y:Q")

# plot bootstrap distribution 
boot_est_dist
```

```python
# plot sampling distribution
sampling_dist
```

Reminder: the true population quantity we are trying to estimate, the population mean, is about 79 years. We know this because we created this population and calculated this value. In real life we wouldn't know this value.

**Question 2.5** True/False
<br> {points: 1}

The mean of the distribution of the bootstrap sample means is the same value as the mean of the sampling distribution of the sample means. True or false?

*Assign your answer to an object called `answer2_5`. Your answer should be a boolean. i.e. `True` or `False`.*

<!-- grade_id: cell-0e3ed3ca2495a0bc, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: cell-3ea212ad593b9f57, points: 1, grade: true, solution: false -->
```python
from hashlib import sha1
assert sha1(str(type(answer2_5)).encode("utf-8")+b"115ea4eaa66205b4").hexdigest() == "b4dbd40b93eeeaef8fae2d611ad0eeb922be4485", "type of answer2_5 is not bool. answer2_5 should be a bool"
assert sha1(str(answer2_5).encode("utf-8")+b"115ea4eaa66205b4").hexdigest() == "4f688fa94471089f9f3bec82911243a7fb93f138", "boolean value of answer2_5 is not correct"

print('Success!')
```

**Question 2.6** True/False
<br> {points: 1}

The mean of the bootstrap sampling distribution is not the same value as the mean of the sampling distribution because the bootstrap sampling distribution was created from samples drawn from a single sample, whereas the sampling distribution was created from samples drawn from the population. True or false?

*Assign your answer to an object called `answer2_6`. Your answer should be a boolean. i.e. `True` or `False`.*

<!-- grade_id: cell-230cf28d4590a63d, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: cell-b167494acba2816d, points: 1, grade: true, solution: false -->
```python
from hashlib import sha1
assert sha1(str(type(answer2_6)).encode("utf-8")+b"e0f0e2aaa82af76a").hexdigest() == "0d31d4fbdd23909e51c3eec22faee488fff0b228", "type of answer2_6 is not bool. answer2_6 should be a bool"
assert sha1(str(answer2_6).encode("utf-8")+b"e0f0e2aaa82af76a").hexdigest() == "e5585473c317ca1a556e01dfb7d8a730aa1d45fd", "boolean value of answer2_6 is not correct"

print('Success!')
```

**Question 2.7** True/False
<br> {points: 1}

The shape and spread (i.e. width) of the distribution of the bootstrap sample means is a poor approximation of the shape and spread of the sampling distribution of the sample means. True or false?

*Assign your answer to an object called `answer2_7`. Your answer should be a boolean. i.e. `True` or `False`.*

<!-- grade_id: cell-f82d46022821b073, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: cell-2ab1350547ae451d, points: 1, grade: true, solution: false -->
```python
from hashlib import sha1
assert sha1(str(type(answer2_7)).encode("utf-8")+b"468ffb90d044c874").hexdigest() == "95ff3ef7d0f1c8d67ad49152b1e0deaebb94ac92", "type of answer2_7 is not bool. answer2_7 should be a bool"
assert sha1(str(answer2_7).encode("utf-8")+b"468ffb90d044c874").hexdigest() == "6321eedbc6e56a84f3f8f98ce9f64f50e6465b61", "boolean value of answer2_7 is not correct"

print('Success!')
```

**Question 2.8** True/False
<br> {points: 1}

In real life, where we only have one sample and cannot create a sampling distribution, the distribution of the bootstrap sample estimates (here means) can suggest how we might expect our point estimate to behave if we took another sample. True or false?

*Assign your answer to an object called `answer2_8`. Your answer should be a boolean. i.e. `True` or `False`.*

<!-- grade_id: cell-59cfd94899e457a7, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: cell-4d2a0325f4fd2406, points: 1, grade: true, solution: false -->
```python
from hashlib import sha1
assert sha1(str(type(answer2_8)).encode("utf-8")+b"5d95bde10dcb6686").hexdigest() == "f41eeb62569e8e334f8a95f7b82f618bc7aa9a8c", "type of answer2_8 is not bool. answer2_8 should be a bool"
assert sha1(str(answer2_8).encode("utf-8")+b"5d95bde10dcb6686").hexdigest() == "b4b927833cb8f695c4bacd36ec9829b008fd23a4", "boolean value of answer2_8 is not correct"

print('Success!')
```

### Using the bootstrap distribution to calculate a plausible range for point estimates

Once we have created a bootstrap distribution, we can use it to suggest a plausible range where we might expect the true population quantity to lie. One formal name for a commonly used plausible range is called a confidence interval. Confidence intervals can be set at different levels, an example of a commonly used level is 95%. When we report a point estimate with a 95% confidence interval as the plausible range, formally we are saying that if we repeated this process of building confidence intervals more times with more samples, wed expect ~ 95% of them to contain the value of the population quantity.

> How do you choose a level for a confidence interval? You have to consider the downstream application of your estimation and what the cost/consequence of an incorrect estimate would be. The higher the cost/consequence, the higher a confidence level you would want to use. You will learn more about this in later Statistics courses.

To calculate an approximate 95% confidence interval using bootstrapping, we essentially order the values in our bootstrap distribution and then take the value at the 2.5th percentile as the lower bound of the plausible range, and the 97.5th percentile as the upper bound of the plausible range. 

```python
# run this cell
np.percentile(boot1000_means["mean"], [2.5, 97.5])
```

Thus, to finish our estimation of the population quantity that we are trying to estimate, we would report the point estimate and the lower and upper bounds of our confidence interval. We would say something like this:

Our sample mean age for Canadian seniors was measured to be 83.7 years, and were 95% confident that the true population mean for Canadian seniors is between (78, 89). 

Here our 95% confidence interval does contain the true population mean for Canadian seniors, 79 years - pretty neat! However, in real life we would never be able to know this because we only have observations from a single sample, not the whole population.

**Question 2.9** True/False
<br> {points: 1}

For any sample we take, if we use bootstrapping to calculate the 95% confidence intervals, the true population quantity we are trying to estimate would always fall within the lower and upper bounds of the confidence interval. True or false?

*Assign your answer to an object called `answer2_9`. Your answer should be a boolean. i.e. `True` or `False`.*

<!-- grade_id: cell-91e462917a53edad, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: cell-e0a24858ef65cea9, points: 1, grade: true, solution: false -->
```python
from hashlib import sha1
assert sha1(str(type(answer2_9)).encode("utf-8")+b"b0039dbcad9a3cc1").hexdigest() == "dc9d76ab0e2e78d7bfd2440330011d7718f1ef32", "type of answer2_9 is not bool. answer2_9 should be a bool"
assert sha1(str(answer2_9).encode("utf-8")+b"b0039dbcad9a3cc1").hexdigest() == "28b8caa2d490dcb9a5bae2428ccbcd15d870910b", "boolean value of answer2_9 is not correct"

print('Success!')
```

# Optional: getting Python + Jupyter working for you outside of our course

At some point after the course is done, you will lose access to the JupyterHub server where you have been doing your course work. If you want to continue to use Python + Jupyter (for other courses at UBC, or for your work after UBC) you have two options:

1. a server solution
2. a local installation solution

We will point you to how you can do both, as well as provide [guidance to take a copy of your homework from our Canvas JupyterHub server](#Getting-your-files-off-of-the-Canvas-JupyterHub).

## Getting your files off of the Canvas JupyterHub

1. On the Canvas JupyterHub open a terminal and type: `tar -zcvf dsci-100.tar.gz dsci-100`
2. Go to the JupyterHub Control Panel/Home and click the box beside the file `dsci-100.tar.gz`, and then click "Download"
3. Log onto either the https://ubc.syzygy.ca/ (CWL authentication) or https://cybera.syzygy.ca/ (Google authentication) servers
4. Click the white "Upload" button (beside the "New" button), and select `tar -zxvf dsci-100.tar.gz`. Then click the blue "Upload" button to complete the upload.
5. open a terminal and type: `tar -zxvf dsci-100.tar.gz` 

You should have your files now available on either of those servers above. 

*On Mac or Linux `tar` also works in terminal, and so you can use this strategy to have the files available locally, however you will need Jupyter installed to view them.*

### 1. A Server Solution

- As a student at UBC, you have access to another JupyterHub that you can access using your UBC CWL: https://ubc.syzygy.ca/
- If you have a Google account, you have access to another JupyterHub that does not depend on you being a UBC student (i.e., having a valid CWL): https://cybera.syzygy.ca/

### 2. A Local Installation Solution

- Depending on your device, you may be able to install Jupyter and Python on it. Below we provide instructions for the 3 major operating systems ([Mac](#A.-Install-Jupyter-+-Python-on-Mac), [Windows](#B.-Install-Jupyter-+-Python-on-Windows) and [Linux]())

### A. Install Jupyter + Python on Mac

#### A.1. Install Python

Go to https://repo.anaconda.com/miniconda/Miniconda3-latest-MacOSX-x86_64.pkg and download the latest version of Python for Mac. Open the file and follow the installer instructions.

#### A.2. Install Jupyter
Head to https://www.anaconda.com/download/#macos and download the Anaconda version for Mac OS. Follow the instructions on that page to run the installer. *Anaconda is a software bundle installs Jupyter and Python and a few other tools. We recommend using it because it makes installation much simpler than other methods.*

#### A.3 Test if it works:

- Open terminal and type: `jupyter notebook`

The above command should open a web browser, with Jupyter's home inside it. Try creating a new Python notebook and running some simple Python code (e.g., `print("hello!")`) to test that it works.

### B. Install Jupyter + Python on Windows

#### B.1. Install Python

Go to https://repo.anaconda.com/miniconda/Miniconda3-latest-Windows-x86_64.exe and download the latest version of Python for Windows. Open the file and follow the installer instructions.

#### B.2. Install Jupyter
Head to <https://www.anaconda.com/download/#windows> and download the Anaconda version for Windows with Python 3.6. After the download has finished, run the installer selecting the following options:
  - On the **Advanced Installation Options** page, check both boxes (see image below)
  - For all other pages, use the default options

<img align="left" src="https://github.com/UBC-MDS/UBC-MDS.github.io/raw/master/resources_pages/imgs/anaconda_windows.png" width="550">


#### B.4 Test if it works:

- Open Windows Command Prompt (select the Start button, type cmd, and click Command Prompt from the list) and type: `jupyter notebook`

The above command should open a web browser, with Jupyter's home inside it. Try creating a new Python notebook and running some simple Python code (e.g., `print("hello!")`) to test that it works.

### C. Install Jupyter + Python on Linux

#### C.1. Install Python

Go to https://repo.anaconda.com/miniconda/Miniconda3-latest-Linux-x86_64.sh and download the latest version of Python for Linux.


#### C.2. Install Jupyter
Head to https://www.anaconda.com/download/#linux and download the Anaconda version for Linux with Python 3.6. Follow the instructions on that page to run the installer. *Anaconda is a software bundle installs Jupyter and Python and a few other tools. We recommend using it because it makes installation much simpler than other methods.*

#### C.3 Test if it works:

- Open terminal and type: `jupyter notebook`

The above command should open a web browser, with Jupyter's home inside it. Try creating a new Python notebook and running some simple Python code (e.g., `print("hello!")`) to test that it works.